Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

For this problem set, we'll be using the Jupyter notebook:

![](jupyter.png)

---
## Part A (2 points)

Write a function that returns a list of numbers, such that $x_i=i^2$, for $1\leq i \leq n$. Make sure it handles the case where $n<1$ by raising a `ValueError`.

<!-- grade_id: squares, points: 0, grade: false, solution: true -->
```python
def squares(n):
    """Compute the squares of numbers from 1 to n, such that the 
    ith element of the returned list equals i^2.
    """
    # YOUR CODE HERE
    list = []
    if n >= 1:
        for i in range(1,n+1):
            list.append(i**2)
    if n < 1:
        raise ValueError('n has to be greater than or equal to zero')
    return list
    
    raise NotImplementedError()
```

Your function should print `[1, 4, 9, 16, 25, 36, 49, 64, 81, 100]` for $n=10$. Check that it does:

```python
squares(10)
```

Output:
```
[1, 4, 9, 16, 25, 36, 49, 64, 81, 100]
```

<!-- grade_id: correct_squares, points: 1, grade: true, solution: false -->
```python
"""Check that squares returns the correct output for several inputs"""
assert squares(1) == [1]
assert squares(2) == [1, 4]
assert squares(10) == [1, 4, 9, 16, 25, 36, 49, 64, 81, 100]
assert squares(11) == [1, 4, 9, 16, 25, 36, 49, 64, 81, 100, 121]
```

<!-- grade_id: squares_invalid_input, points: 1, grade: true, solution: false -->
```python
"""Check that squares raises an error for invalid inputs"""
try:
    squares(0)
except ValueError:
    pass
else:
    raise AssertionError("did not raise")

try:
    squares(-4)
except ValueError:
    pass
else:
    raise AssertionError("did not raise")
```

---

## Part B (1 point)

Using your `squares` function, write a function that computes the sum of the squares of the numbers from 1 to $n$. Your function should call the `squares` function -- it should NOT reimplement its functionality.

<!-- grade_id: sum_of_squares, points: 0, grade: false, solution: true -->
```python
def sum_of_squares(n):
    """Compute the sum of the squares of numbers from 1 to n."""
    # YOUR CODE HERE
    tot = 0
    a = squares(n)
    for i in a:
        tot += i
    return tot
    raise NotImplementedError()
```

The sum of squares from 1 to 10 should be 385. Verify that this is the answer you get:

```python
sum_of_squares(10)
```

Output:
```
385
```

<!-- grade_id: correct_sum_of_squares, points: 0.5, grade: true, solution: false -->
```python
"""Check that sum_of_squares returns the correct answer for various inputs."""
assert sum_of_squares(1) == 1
assert sum_of_squares(2) == 5
assert sum_of_squares(10) == 385
assert sum_of_squares(11) == 506
```

<!-- grade_id: sum_of_squares_uses_squares, points: 0.5, grade: true, solution: false -->
```python
"""Check that sum_of_squares relies on squares."""
orig_squares = squares
del squares
try:
    sum_of_squares(1)
except NameError:
    pass
else:
    raise AssertionError("sum_of_squares does not use squares")
finally:
    squares = orig_squares
```

---
## Part C (1 point)

Using LaTeX math notation, write out the equation that is implemented by your `sum_of_squares` function.

<!-- grade_id: sum_of_squares_equation, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

$$ \sum_{i=1}^{n}x^2_i $$

---
## Part D (2 points)

Find a usecase for your `sum_of_squares` function and implement that usecase in the cell below.

$$\sum_{i=1}^{n}(x_i - \bar{x})^2 =  \sum_{i=1}^{n}x^2_i - \frac{1}{n}\left [ \sum_{i=1}^{n}x_i \right ]^2 $$


<!-- grade_id: sum_of_squares_application, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
# To calculate the sum of squares
# Sum of square
n = 5
a = sum_of_squares(n)

# Sum of x_i square
def sum_of_num_square(n):
    tot = 0
    if n >= 1:
        for i in range(1,n+1):
            tot += i
    else:
        raise ValueError('n has to be greater than or equal to zero')
    tot_square = tot**2
    return tot_square
    raise NotImplementedError()

b = sum_of_num_square(n)
tot_sum_of_square = a - (1/n)*b
print(tot_sum_of_square)
```

Output:
```
10.0
```

---
## Part E (4 points)

State the formulae for an arithmetic and geometric sum and verify them numerically for an example of your choice.

#### Arithmetic sum:

$$ S_n = \frac{n}{2}(2a_1 + (n-1)d) $$

```python
def arithmetic_sum(start, n, diff):
    sum = 0
    if n % 1 == 0 and n >= 1:
        sum = (n/2)*(2*start + (n-1)*diff)
    else:
        raise ValueError("Give an integer and/or a value greater than 0")
    return sum
print(arithmetic_sum(1, 10, 6))
```

Output:
```
280.0
```

#### Geometric sum:

$$ S_n = \frac{a_1(1 -  r^{n})}{1-r} $$

```python
def geom_sum(start, n, ratio):
    sum = 0
    if n % 1 == 0 and n >= 1:
        sum = (start*(1-ratio**n))/(1-ratio)
    else:
        raise ValueError("Give an integer and/or a value greater than 0")
    return sum
print(geom_sum(1, 10, 6))
```

Output:
```
12093235.0
```